**Before you turn this problem in, make sure your notebook is in tidy format and everything runs as expected:**

1. When the problem needs you to write codes inside a function, just directly write codes inside the already-existing functions. Don't create new functions and write your code in the new functions. 

2. Don't modify the original TEST cells! You can insert new testing cells to test your code when doing your homework. But make sure to delete those **new testing cells** after you finish your problem.

3. After you complete all problems, to insure there is no error in your notebook, do the following: first, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and second **run all cells** (in the menubar, select Cell$\rightarrow$Run All). **Note that this step is very important and can help you catch the error!**

**Make sure you delete `raise NotImplementedError()`** (if existing) inside the function before you write your code and fill in any place that says "YOUR CODE HERE" or "YOUR ANSWER HERE":

```python
# fill in your name
NAME = ""
```

---

For many of the coding assignments we give you, we will provide tests that your code must pass. These will be in one or more cells after the code you need to write, and they *must* pass for you to get credit.

Also, be aware that test cells will be included only for *coding* questions -- you're on your own to judge whether you have the correct answer for questions requiring a written response!

**This file is only for testing purpose, and it is not homework! Don't submit this file to canvas.**

## 1. (1 point)

Markdown is a special way of writing text in order to specify formatting, like whether text should be bold, italicized, etc.

For the following few exercises, try to format the specified text using Markdown.


* Hint: after editing the Markdown, you will need to run the cell so that the formatting appears.

Note that in this assignment, the formatted text boxes are actually images, so you cannot look at the cell in edit mode to see the answer. However, in future assignments, we will write all the instructions in Markdown, so in the future you will always be able to look at the cell in edit mode to see how the formatting is done.

<div class="alert alert-success">
Italicize the word "Banana", like this:
</div>

![banana](images/banana-italics.png)

<!-- grade_id: test_markdown, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## 2. (2 points)

<div class="alert alert-success">
Implement the function <code>hello</code> and make sure the test cells runs without any errors. You will need to delete the line with `raise NotImplementedError`, write your own solution, and then re-run the cell before running the test cell. Each time you change code in a cell, you  will need to re-run that cell before running any other cells that depend on it!
</div>

```python
from nose.tools import assert_equal
```

<!-- grade_id: hello, points: 0, grade: false, solution: true -->
```python
def hello(name):
    """Returns a message containing "Hello, <name>!", 
    where <name> is passed in as an argument.
    
    Parameters
    ----------
    name : string
        The name of the person to say hello to
        
    Returns
    -------
    the message containing "Hello, <name>!"
    
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: test_hello, points: 0, grade: true, solution: false -->
```python
# this is the formal TEST cell and you should not modify!!!

"""(1 point) Test code for the 'hello' function. This cell should NOT give any errors when it is run."""
from nose.tools import assert_equal
assert_equal(hello("Jessica"), "Hello, Jessica!")
assert_equal(hello("jessica"), "Hello, jessica!")
assert_equal(hello("Tom"), "Hello, Tom!")
assert_equal(hello("123"), "Hello, 123!")

print("Success!")
```

<div class="alert alert-info">When completing a problem set, it is often useful to supplement the autograder tests with test cases of your own. Below, we provide you with a "scratch" cell to be used for your own testing purposes. <br>

Keep in mind, however, that writing your own test cases is **completely optional** -- these are only here to help you test whether your answer is correct as you complete the assignment. This means that you will *not* be graded on anything you include in these cells you created yourself!</div>

```python
# add your own test cases in this cell!

# try running your hello function with your own name and see what happens
# this is to compare whether the output from hello("shuang") is the same as the string "Hello, shuang!"

assert_equal(hello("shuang"), "Hello, shuang!")
```